```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# Lab 01

**Welcome to the first lab of DATA 2201!** This lab is meant to help you familiarize yourself with JupyterHub, review Python and `NumPy`, and introduce you to `matplotlib`, a Python visualization library. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

You must submit this assignment to Gradescope by the on-time deadline

**We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to contact staff for submission support. 

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 0: Jupyter Notebooks - Review

This file `lab01.ipynb` is called a Jupyter Notebook.  

###  Text (Markdown) cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You don't need to learn Markdown now, but you might want to.

After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. (Try not to delete the instructions of the lab.)

### Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

Try running this cell:

```python
print("Hello, World")
```

###  Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a text cell.  You can change it to a code cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart kerneal button in the menu bar, and choosing "Code".

**Check your understanding** Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell! ♪🌏♪

(That musical note symbol is like the Earth symbol.  Its long-form name is `\N{EIGHTH NOTE}`.)

Run your cell to verify that it works.

### The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
3. If that doesn't help, restart your server. First, save your work by clicking **File** at the top left of your screen, then **Save Notebook**. Next, click **File** then **Hub Control Panel**. In the new windown, choose **Stop My Server** to shut it down, then **Start My Server** to start it back up. Then, navigate back to the notebook you were working on. You'll still have to run your code cells again.

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 1: Jupyter Tips

### Viewing Documentation

To output the documentation for a function, use the `help` function.

```python
help(print)
```

You can also use Jupyter to view function documentation inside your notebook. The function must already be defined in the kernel for this to work.

Below, click your mouse anywhere on the `print` block below and use `Shift` + `Tab` to view the function's documentation. 

```python
print('Welcome to DATA 2201.')
```

### Importing Libraries and Magic Commands

In DATA 2201, we will be using common Python libraries to help us process data. By convention, we import all libraries at the very top of the notebook. There are also a set of standard aliases that are used to shorten the library names. Below are some of the libraries that you may encounter throughout the course, along with their respective aliases.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
%matplotlib inline
```

`%matplotlib inline` is a [Jupyter magic command](https://ipython.readthedocs.io/en/stable/interactive/magics.html) that configures the notebook so that `matplotlib` displays any plots that you draw directly in the notebook rather than to a file, allowing you to view the plots upon executing your code. (Note: In practice, this is no longer necessary, but we're showing it to you now anyway.)

Another useful magic command is `%%time`, which times the execution of that cell. You can use this by writing it as the first line of a cell. (Note that `%%` is used for *cell magic commands* that apply to the entire cell, whereas `%` is used for *line magic commands* that only apply to a single line.)

```python
%%time
lst = []
for i in range(100):
    lst.append(i)
```

### Keyboard Shortcuts

Even if you are familiar with Jupyter, we strongly encourage you to become proficient with keyboard shortcuts (this will save you time in the future). To learn about keyboard shortcuts, go to **Help --> Keyboard Shortcuts** in the menu above. 

Here are a few that we like:
1. `Ctrl` + `Return` (or `Cmd` + `Return` on Mac): *Evaluate the current cell*
2. `Shift` + `Return`: *Evaluate the current cell and move to the next*
3. `Ctrl` + `+` + `/`: *Comment or uncomment the selected code at once*
4. `ESC` : *command mode* (may need to press before using any of the commands below)
    1. `a` : *create a cell above*
    2. `b` : *create a cell below*
    3. `dd` : *delete a cell*
    4. `z` : *undo the last cell operation*
    5. `m` : *convert a cell to markdown*
    6. `y` : *convert a cell to code*

### Running Cells

Aside from keyboard shortcuts (specifically `Shift` + `Return`), you can also run a single cell by clicking the **Run** button in the top left corner of your notebook. If you hover over the button, you will also find some other options that allow you to run multiple cells. Specifically, the **Run All Above Selected Cell** option is particularly useful for situations wherein you have restarted your notebook and need to run all the cells up until the question you were working on in a lab/homework. 

<br><br>

<hr style="border: 1px solid #ffcd00;" />

## Part 2: Prerequisites

It's time to answer some review questions. Each question has a response cell directly below it. Most response cells are followed by a test cell that runs automated tests to check your work. Please don't delete questions, response cells, or test cells. You won't get credit for your work if you do.

If you have extra content in a response cell, such as an example call to a function you're implementing, that's fine. Also, feel free to add cells between the question cells and test cells (or the next cell, for questions without test cases). Any extra cells you add will be considered part of your submission. Finally, when you finish an assignment, make sure to "restart and run all cells" to ensure everything works properly. 

Note that for labs, on-time submissions that pass all the test cases will receive full credit. However, for homeworks, test cells don't always confirm that your response is correct. They are meant to give you some useful feedback, but it's your responsibility to ensure your response answers the question correctly. There may be other tests that we run when scoring your notebooks. We **strongly recommend** that you check your solutions yourself rather than just relying on the test cells.

### Python

Python is the main programming language we'll use in the course. We expect that you've taken DATA 1202 or CS 2321, so you have either had a course introducing Python through data science or have taken several programming courses. This course will not be focusing on general Python syntax. If any of the following exercises are challenging (or if you would like to learn more about Python), please review one or more of the following materials.

- **[Python Tutorial](https://docs.python.org/release/3.10.0/tutorial/index.html)**: Introduction to Python from the creators of Python.
- **[Composing Programs Chapter 1](https://www.composingprograms.com/pages/11-getting-started.html)**: This is more of an introduction to programming with Python.
- **[Advanced Crash Course](http://cs231n.github.io/python-numpy-tutorial/)**: A fast crash course which assumes some programming background.

### `NumPy`

`NumPy` is the numerical computing module. Here's a quick recap of `NumPy`. For more review, read the following materials.

- **[NumPy Quick Start Tutorial](https://numpy.org/doc/stable/user/quickstart.html)**
- **[NumPy Review](https://mtu.instructure.com/courses/1527249/files/118524356?module_item_id=24078621&fd_cookie_set=1)**
- **[Stanford CS231n NumPy Tutorial](http://cs231n.github.io/python-numpy-tutorial/#numpy)**
- **[The Data 8 Textbook Chapter on NumPy](https://www.inferentialthinking.com/chapters/05/1/Arrays)**

<br/><br/>

---

#### Question 1

The core of `NumPy` is the array. Like Python lists, arrays store data; however, they store data in a more efficient manner. In many cases, this allows for faster computation and data manipulation.

In DATA 1202, we used `make_array` from the `datascience` module, but that's not the most typical way. Instead, use `np.array` to create an array. It takes a sequence, such as a list or range. 

Below, create an array `arr` containing the values 1, 2, 3, 4, and 5 (in that order).

```python
arr = ...
arr
```

```python
grader.check("q1")
```

In addition to values in the array, we can access attributes such as shape and data type. A full list of attributes can be found [here](https://docs.scipy.org/doc/numpy-1.15.0/reference/arrays.ndarray.html#array-attributes).

```python
arr[3]
```

```python
arr[4:5]
```

```python
arr[2:4]
```

```python
arr.shape
```

```python
arr.dtype
```

Arrays, unlike Python lists, cannot store items of different data types.

```python
# A regular Python list can store items of different data types
[1, '3']
```

```python
# Arrays will convert everything to the same data type
np.array([1, '3'])
```

```python
# Another example of array type conversion
np.array([5, 8.3])
```

Arrays are also useful in performing __*vectorized operations*__. Given two or more arrays of equal length, arithmetic will perform element-wise computations across the arrays. 

For example, observe the following:

```python
# Python list addition will concatenate the two lists
[1, 2, 3] + [4, 5, 6]
```

```python
# NumPy array addition will add them element-wise
np.array([1, 2, 3]) + np.array([4, 5, 6])
```

<br><br>

<hr style="border: 1px solid #ffcd00;" />

## Part 3: Summations and LaTeX

You should use LaTeX to format math in your answers. If you aren't familiar with LaTeX, don't worry. It's not hard to use in a Jupyter notebook. Just place your math in between dollar signs within Markdown cells:

`$ f(x) = 2x $` becomes $ f(x) = 2x $.

If you have a longer equation, use double dollar signs to place it on a line by itself:

`$$ \sum_{i=0}^n i^2 $$` becomes:

$$ \sum_{i=0}^n i^2$$


You can align multiple lines using the `&` anchor, `\\` newline, in an `align` block as follows:

```
\begin{align}
f(x) &= (x - 1)^2 \\
&= x^2 - 2x + 1
\end{align}
```
becomes

\begin{align}
f(x) &= (x - 1)^2 \\
&= x^2 - 2x + 1
\end{align}

* [This PDF](https://mtu.instructure.com/courses/1527249/files/118523182?module_item_id=24078678) has some handy LaTeX tips.
* [For more about basic LaTeX formatting, you can read this article.](https://www.overleaf.com/learn/latex/Mathematical_expressions)


<br/>

---

### Question 2

Write a function `summation` that evaluates the following summation for $n \geq 1$:

$$\sum_{i=1}^{n} i^3 + 3 i^2$$

**Note**: You should not use `for` loops in your solution. Check the [NumPy documentation](https://numpy.org/doc/stable/reference/index.html). If you're stuck, try a search engine! Searching the web for examples of how to use modules is very common in data science. You may find `np.arange`, `np.power`, and other such funcitons helpful for this question!

```python
def summation(n):
    """Compute the summation i^3 + 3 * i^2 for 1 <= i <= n."""
    ...
```

```python
grader.check("q2")
```

<br/>

---

### Question 3

Write a function `elementwise_array_sum` that computes the square of each value in `list_1`, the cube of each value in `list_2`, then returns a list containing the element-wise sum of these results. Assume that `list_1` and `list_2` have the same number of elements, do not use for loops.

The input parameters will both be **Python lists**, so you may need to convert the lists into arrays before performing your operations. The output should be a **`NumPy` array.**

```python
def elementwise_array_sum(list_1, list_2):
    """Compute x^2 + y^3 for each x, y in list_1, list_2. 
    
    Assume list_1 and list_2 have the same length.
    
    Return a NumPy array.
    """
    assert len(list_1) == len(list_2), "both args must have the same number of elements"
    ...
```

```python
grader.check("q3")
```

You might have been told that Python is slow, but array arithmetic is carried out very fast, even for large arrays. Below is an implementation of the above code that does not use `NumPy` arrays.

```python
def elementwise_list_sum(list_1, list_2):
    """Compute x^2 + y^3 for each x, y in list_1, list_2. 
    
    Assume list_1 and list_2 have the same length.
    """

    return [x ** 2 + y ** 3 for x, y in zip(list_1, list_2)]
```

For ten numbers, `elementwise_list_sum` and `elementwise_array_sum` both take a similar amount of time.

```python
sample_list_1 = list(range(10))
sample_array_1 = np.arange(10)
```

```python
%%time
elementwise_list_sum(sample_list_1, sample_list_1)
```

```python
%%time
elementwise_array_sum(sample_array_1, sample_array_1)
```

The time difference seems negligible for a list/array of size 10; depending on your setup, you may even observe that `elementwise_list_sum` executes faster than `elementwise_array_sum`! However, we will commonly be working with much larger datasets:

```python
sample_list_2 = list(range(100000))
sample_array_2 = np.arange(100000)
```

```python
%%time
elementwise_list_sum(sample_list_2, sample_list_2)
# The semicolon hides the output
; 
```

```python
%%time
elementwise_array_sum(sample_array_2, sample_array_2)
# The semicolon hides the output
; 
```

With the larger dataset, we see that using `NumPy` results in code that executes over 20 times faster! Throughout this course (and in the real world), you will find that writing efficient code will be important; arrays and vectorized operations are the most common way of making Python programs run quickly.

<br/>

---

### Question 4
In statistics, the formula for population variance is:

$$\sigma^2 = \frac{\sum_{i=1}^N (x_i - \mu)^2}{N}$$

The variance is the mean squared deviation from average, $\mu$. 

Complete the functions below to compute the population variance of `population`, an array of numbers. For this question, **do not use built-in  `NumPy ` functions, such as `np.var`.** Again, avoid using `for` loops! For a refresher on what variance is, feel free to read up on it in the Textbook [here](https://inferentialthinking.com/chapters/14/2/Variability.html)!

```python
def mean1(population):
    """
    Returns the mean of population (mu)
    
    Keyword arguments:
    population -- a numpy array of numbers
    """
    # Calculate the mean of a population
    ...

def variance1(population):
    """
    Returns the variance of population (sigma squared)
    
    Keyword arguments:
    population -- a numpy array of numbers
    """
    # Calculate the variance of a population
    ...
```

```python
grader.check("q4")
```

<br/>

---

### Question 5

Given the array `random_arr`, assign `valid_values` to an array containing all values $x$ such that $2x^4 > 1$.

**Note**: You should not use `for` loops in your solution. Instead, look at  `NumPy`'s documentation on [Boolean Indexing](https://numpy.org/doc/stable/user/basics.indexing.html#boolean-array-indexing). Documentation can be very intimidating at first glance, but don't worry, that's completely okay, one of the goals of this class is to build familiarity with reading the documentation of data science tools. Ask for help if needed, we're always there for you!

```python
np.random.seed(42)
random_arr = np.random.rand(60)
valid_values = ...
valid_values
```

```python
grader.check("q5")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Congratulations! You have finished Lab 01! 

## Submission 

Make sure you have run all cells in your notebook in order before running the cell below. The cell below will generate a zip file for you to submit. 

**Please save before exporting!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: [2, 2])
```
>>> type(arr) == np.ndarray
True
```
```
>>> all(arr == np.array([1, 2, 3, 4, 5]))
True
```


Tests for q2 (points: [2, 2, 2, 2])
```
>>> summation(1) == 4
np.True_
```
```
>>> summation(2) == 24
np.True_
```
```
>>> summation(10) == 4180
np.True_
```
```
>>> summation(100) == 26517550
np.True_
```


Tests for q3 (points: 10)
```
>>> ae = False
>>> try:
...     elementwise_array_sum([1, 2, 3], [1, 2, 3, 4])
... except AssertionError:
...     ae = True
>>> ae == True
True
```
```
>>> all(elementwise_array_sum([], []) == np.array([]))
True
```
```
>>> all(elementwise_array_sum([1], [1]) == np.array([2]))
True
```
```
>>> all(elementwise_array_sum([-1], [1]) == np.array([2]))
True
```
```
>>> all(elementwise_array_sum([1], [-1]) == np.array([0]))
True
```
```
>>> all(elementwise_array_sum([1, 2, 3], [1, 2, 3]) == np.array([2, 12, 36]))
True
```
```
>>> all(elementwise_array_sum([1, 5, 2], [3, 6, 6]) == np.array([28, 241, 220]))
True
```
```
>>> all(elementwise_array_sum([1, 1], [1, 1]) == np.array([2, 2]))
True
```
```
>>> all(np.isclose(elementwise_array_sum([1.5, 1], [1, 1]), [3.25, 2.0], 0.01))
True
```
```
>>> type(elementwise_array_sum([], [])) == np.ndarray
True
```


Tests for q4 (points: 10)
```
>>> population_0 = np.random.randn(100)
>>> np.isclose(mean1(population_0), np.mean(population_0), 1e-06)
np.True_
```
```
>>> population_1 = 3 * np.random.randn(100) + 5
>>> np.isclose(mean1(population_1), np.mean(population_1), atol=1e-06)
np.True_
```
```
>>> population_0 = np.random.randn(100)
>>> np.isclose(variance1(population_0), np.var(population_0), atol=1e-06)
np.True_
```
```
>>> population_1 = 3 * np.random.randn(100) + 5
>>> np.isclose(variance1(population_1), np.var(population_1), atol=1e-06)
np.True_
```


Tests for q5 (points: 10)
```
>>> len(valid_values) == 10
True
```
```
>>> np.allclose(valid_values, [0.95071431, 0.86617615, 0.96990985, 0.94888554, 0.96563203, 0.9093204, 0.96958463, 0.93949894, 0.89482735, 0.92187424])
True
```
